Growth in early infancy drives optimal brain functional connectivity which predicts cognitive flexibility in later childhood

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

This study utilises fNIRS to investigate the effects of undernutrition on functional connectivity patterns in infants from a rural population in Gambia. fNIRS resting-state data recording spanned ages 5 to 24 months, while growth measures were collected from birth to 24 months. Additionally, executive functioning tasks were administered at 3 or 5 years of age. The results show an increase in left and right frontal-middle and right frontal-posterior connections with age and, contrary to previous findings in high-income countries, a decrease in frontal interhemispheric connectivity. Restricted growth during the first months of life was associated with stronger frontal interhemispheric connectivity and weaker right frontal-posterior connectivity at 24 months of age. Additionally, the study describes some connectivity patterns, including stronger frontal interhemispheric connectivity, which is associated with better cognitive flexibility at preschool age.

Strengths:

The study analyses longitudinal data from a large cohort (n = 204) of infants living in a rural area of Gambia. This already represents a large sample for most infant studies, and it is impressive, considering it was collected outside the lab in a population that is underrepresented in the literature. The research question regarding the effect of early nutritional deficiency on brain development is highly relevant and may highlight the importance of early interventions. The study may also encourage further research on different underrepresented infant populations (i.e., infants not residing in Western high-income countries) or in settings where fMRI is not feasible.

The preprocessing and analysis steps are carefully described, which is very welcome in the fNIRS field, where well-defined standards for preprocessing and analysis are still lacking.

We thank the reviewer for highlighting the strengths of this work.

Weaknesses:

While the study provides a solid description of the functional connectivity changes in the first two years of life at the group level and investigates how restricted growth influences connectivity patterns at 24 months, it does not explore the links between adverse situations and developmental trajectories for functional connectivity. Considering the longitudinal nature of the dataset, it would have been interesting to apply more sophisticated analytical tools to link undernutrition to specific developmental trajectories in functional connectivity. The authors mention that they lack the statistical power to separate infants into groups according to their growing profiles. However, I wonder if this aspect could not have been better explored using other modelling strategies and dimensional reduction techniques. I can think about methods such as partial least squares correlation, with age included as a numerical variable and measures of undernutrition.

We agree with the reviewer that this complex and rich longitudinal dataset would benefit from more sophisticated analytical approaches to characterise developmental trajectories in functional connectivity and to more directly link them to measures of undernutrition. However, conducting such analyses would require substantial additional methodological development, model validation, and careful interpretation, which fall beyond the scope and timeline of the present manuscript. Our aim here was to provide a clear and robust characterisation of functional connectivity changes during the first two years of life and to examine associations with growth outcomes at a specific developmental stage, while ensuring methodological transparency and statistical reliability. Importantly, these more advanced trajectory-based analyses are currently being pursued in the final phase of the BRIGHT project (BRIGHT IMPACT), in collaboration with expert statisticians and data scientists. This ongoing work aims specifically to leverage the longitudinal richness of the dataset to model developmental trajectories and their associations with early-life adversity and nutritional factors. We therefore see the present study as an important foundation for these forthcoming analyses.

Connectivity was assessed in 6 big ROIs. While the authors justify this choice to reduce variability due to head size and optodes placement, this also implies a significant reduction in spatial resolution. Individual digitalisation and co-registration of the optodes to the head model, followed by image reconstruction, could have provided better spatial resolution. This is not a weakness specific to this study but rather a limitation common to most fNIRS studies, which typically analyse data at the channel level since digitalisation and co-registration can be challenging, especially in complex setups like this. However, the BRIGHT project has demonstrated that it is possible and that differences in placement affect activation patterns, which become more localised when data is co-registered at the subject level (Collins-Jones et al., 2021). Could the co-registration of individual data have increased sensitivity, particularly given that longitudinal effects are being investigated?

We agree with the reviewer that the fNIRS community should work toward more precise methods for spatial registration of optodes, not only at the group level but also at the subject level, in order to make more precise inferences about the locations of activations. However, we followed a very thorough offline procedure to model headgear placement based on each participant’s photographs, which we believe complements the coregistration work performed by Collins-Jones in 2021. As reported in the fNIRS data acquisition section “Infants were excluded from further analysis if the band was excessively high over the front above the eyebrows” (line 409, methods section). Moreover channels displacement was measured from the photos, and if it was “equal or greater than 1.6 cm were renumbered, so that each channel was shifted either backward or forward one full channel location in space” (line 413, methods section). While these practices are thoroughly followed in the BRIGHT project, we are aware that they are not part of the standard procedure in many infant fNIRS studies. We hope that this work provides guidance for other researchers on how to coregister infant fNIRS data.

Considering the spatial resolution of fNIRS, which is on the order of centimetres, and the thorough procedure combining fNIRS–MRI coregistration with channel displacement assessment based on photographs, we do not think that individual-level coregistration would have significantly increased the sensitivity of the results.

I believe that a further discussion in the manuscript on the application of global signal regression and its effects could have been beneficial for future research and for readers to better understand the negative correlations described in the results. Since systemic physiological changes affect HbO/HbR concentrations, resulting in an overestimation of functional connectivity, regressing the global signal before connectivity computation is a common strategy in fNIRS and fMRI studies. However, the recommendation for this step remains controversial, likely depending on the case (Murphy & Fox, 2017). I understand that different reasons justify its application in the current study. In addition to systemic physiological changes originating from brain tissue, fNIRS recordings are contaminated by changes occurring in superficial layers (i.e., the scalp and skull). While having short-distance channels could have helped to quantify extracerebral changes, challenges exist in using them in infant populations, especially in a longitudinal study such as the one presented here. The optimal source-detector distance that minimises sensitivity to changes originating from the brain would increase with head size, and very young participants would require significantly shorter source-detector distances (Brigadoi & Cooper, 2015). Thus, having them would have been challenging. Under these circumstances (i.e., lack of short channels and external physiological measures), and considering that the amount the signal is affected by physiological noise (either coming from the brain or superficial tissue) might change through development, the choice of applying global signal regression is justified. Nevertheless, since the method introduces negative correlations in the data by forcing connectivity to average to zero, I believe a further discussion of these points would have enriched the interpretation of the results.

We added a paragraph discussing the choice of using GSR in our pipeline in the discussion of the manuscript as follows: “Importantly, these results remained significant even without GSR, indicating that our findings are not solely driven by preprocessing choices. While the use of GSR in FC studies remains debated (Murphy & Fox, 2017), in the absence of short channels (which are difficult to use reliably with infants (Emberson et al., 2016) and external physiological measures, applying GSR represented the most appropriate preprocessing option. In fact, failure to correct for systemic physiological fluctuations can, in fact, lead to artificially elevated connectivity estimates in fNIRS data (Abdalmalak et al., 2022)” (line 250, discussion section).

Reviewer #2 (Public review):

Strengths:

The article addresses a topic of significant importance, focusing on early life growth faltering in low-income countries-a key marker of undernutrition-and its impact on brain functional connectivity (FC) and cognitive development. The study's strengths include the laborious data collection process, as well as the rigorous data preprocessing methods employed to ensure high data quality. The use of cutting-edge preprocessing techniques further enhances the reliability and validity of the findings, making this a valuable contribution to the field of developmental neuroscience and global health.

We thank the reviewer for highlighting the strengths of this work.

Weaknesses:

The study fails to fully leverage its longitudinal design to explore neurodevelopmental changes or trajectories, as highlighted by all three reviewers. The revised manuscript still primarily focuses on FC values at a single age stage (i.e., 24 months) rather than utilizing the longitudinal data to investigate how FC evolves over time or predicts cognitive development. Although the authors acknowledge that analyzing changes in FC (ΔFC) would reduce degrees of freedom (to ~30) and risk interpretability, they do not report or discuss these results, even as exploratory findings.

As suggested, we added the table reporting the results of the associations between changes in functional connectivity (DFC) between 5 and 24 months and cognitive flexibility in the supplementary materials (Table SI3). We additionally explored the relationship between changes in growth and cognitive flexibility as suggested by Reviewer #3 and we reported these additional analyses in the text as follows: “We also explored whether changes in growth and changes in functional connectivity between 5 and 24 months were associated with cognitive flexibility at preschool age, but we did not find any significant association (Table SI3 and Table SI4).” (line 213, results section).

Furthermore, the study lacks specificity in identifying which specific brain networks are affected by growth faltering, as the current exploratory analyses mainly provide an overall conclusion that infant brain network development is impacted without pinpointing the precise neural mechanisms or networks involved.

We added this limitation in the discussion as follows: “While the impact of undernutrition on brain development has been documented in LMICs (46), herein, we provided empirical evidence that growth faltering specifically in infants younger than five months of age impacts observable development of functional brain networks in the second year of life. Future studies may be needed to pinpoint which specific brain networks are impacted” (line 279, discussion section).

Reviewer #3 (Public review):

Summary

This study aimed to investigate whether the development of functional connectivity (FC) is modulated by early physical growth, and whether these might impact cognitive development in childhood. This question was investigated by studying a large group of infants (N=204) assessed in Gambia with fNIRS at 5 visits between 5 and 24 months of age. Given the complexity of data acquisition at these ages and following data processing, data could be analyzed for 53 to 97 infants per age group. FC was analyzed considering 6 ensembles of brain regions and thus 21 types of connections. Results suggested that: i) compared to previously studied groups, this group of Gambian infants have different FC trajectory, in particular with a change in frontal inter-hemispheric FC with age from positive to null values; ii) early physical growth, measured through weight-for-length z-scores from birth on, is associated with FC at 24 months. Some relationships were further observed between FC during the first two years and cognitive flexibility, in different ways between 4- and 5-year-old preschoolers, but results did not survive corrections for multiple comparisons.

Strengths

The question investigated in this article is important for understanding the role of early growth and undernutrition on brain and behavioral development in infants and children. The longitudinal approach considered is highly relevant to investigate neurodevelopmental trajectories. Furthermore, this study targets a little studied population from a low-/middle-income country, which was made possible by the use of fNIRS outside the lab environment. The collected dataset is thus impressive and it opens up a wide range of analytical possibilities.

We thank the reviewer for highlighting the strengths of this work.

Weaknesses

Data analyses were constrained by the limited number of children with longitudinal data on NIRS functional connectivity. Nevertheless, considering more advanced statistical modelling approaches would be relevant to further explore neurodevelopmental trajectories as well as relationships with early growth and later cognitive development.

While in this study we selected specific FC and outcome variables based on our hypothesis, the final phase of the BRIGHT project, known as BRIGHT IMPACT, aims to apply advanced statistical models to integrate a range of project variables into a single comprehensive analysis. We have acknowledged this in the discussion as follows: “Applying more advanced statistical modelling methods and structural equation modelling analyses may provide greater insight with further investigations in contexts of adversity and, in turn, establish which outcomes are predicted by FC” (line 309, discussion section).

The abstract and end of the discussion should make it clearer that the associations between FC and cognitive flexibility are results that need to be confirmed, insofar as they did not survive correction for multiple comparisons.

We have acknowledged this in the abstract as follows: “Our results highlight the measurable effects that poor growth in early infancy has on brain development and the possible subsequent impact on pre-school age cognitive development, underscoring the need for early life interventions throughout global settings of adversity”.

We have acknowledged this in the discussion as follows: “While our results are consistent with previous studies, we acknowledge that the significant associations between early FC and later cognitive flexibility do not withstand multiple comparisons. Therefore, we encourage future studies that may replicate these findings with a larger sample” (line 300, discussion section).

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

(1) In Figure 1 B and C the authors should indicate that the results refer to HbO.

We have added the suggested specification in the caption of the figure as suggested.

(2) Figure SI2. Please indicate in the caption that these are the results when pre-processing did not include global signal regression.

We have added the suggested specification in the caption of the figure as suggested.

Reviewer #3 (Recommendations for the authors):

(1) The sentence l529-531 ("To investigate whether FC early in life predicted...") should be more explicit as it is not clear which of the two variables is regressed by the other: is it the measure of cognitive flexibility that is regressed by FC, as the hypothesis suggests? Were other variables considered in the regression model? (For linear regression with only one "prediction" variable, the square root of the coefficient of determination 𝑅2 is equal to the correlation between the two variables.)

Yes, it is the measure of cognitive flexibility that is regressed by FC. We have rephrased it in the text as follows: “we regressed later cognitive flexibility against FC that showed a significant change across the first two years of life”. There were no other variables in the regression model.

(2) A summary table of the statistical results for FC-cognitive flexibility associations should be included as for other analyses, in addition to Figure 3B.

We added a table of the results for the association between FC and cognitive flexibility in the supplementary materials (Table SI2, page 10), matching the same colours of Table 2. We referenced the table in the text in the main manuscript (line 211, result section).

(3) Figure 3B: The legend should precise that these results did not survive corrections for multiple comparisons.

We have specified this in the legend of Figure 3 as suggested.

(4) For the young pre-schooler group, it seems that the age is around 4 years (age mean +/- SD=47.96 +/- 2.77 months) and not 3 years as indicated at several places in the manuscript.

We found only once instance in which we erroneously said that the younger preschoolers were around 3 years. We replaced “Gambian infants from BRIGHT were cross-sectionally assessed at the age of 3 or 5 years for cognitive flexibility” with Gambian infants from BRIGHT were cross-sectionally assessed between the age of 3 and 5 years for cognitive flexibility (line 489, method section).

(5) The authors use the term "intra-hemispheric" connections for the ones within each of the 6 sections. This might be misleading since fronto-posterior connections are also intra-hemispheric ones. Specifying "short-range" or "within-section" connections might be clearer.

As suggested by the reviewer, we replaced “intra-hemispheric” with “intra-hemispheric within section” where appropriate through the whole manuscript.

(6) Abstract: what is the justification for using the term "optimal" for describing developmental trajectories of FC?

The term “optimal” refers to knowledge about typical developmental trajectories, coming especially from fMRI studies, as mentioned in the introduction: “Based on data from fMRI, current models hypothesize that FC patterns mature throughout early development (23–27), where in typically developing brains, adult-like networks emerge over the first years of life as long-range functional connections between pre-frontal, parietal, temporal, and occipital regions become stronger and more selective (28–31). [...]. Importantly, normative developmental patterns may be disrupted and even reversed in clinical conditions that impact development; e.g., increased short-range and reduced long-range FC have been observed in preterm infants (36) and in children with autism spectrum disorder (37, 38)” (line 93-106, introduction).

(7) The confidence interval should be added in Figure SI3.

As suggested, confidence intervals have been added in Figure SI3.

(8) Other scatterplot examples of associations might be added as supplementary information.

As suggested, we added several additional scatterplots to Figure SI3 (with confidence intervals as noted in the comment above) to show other associations between changes in growth and FC at 24 months.

(9) Figure SI6: % in x-axis is still indicated.

We apology for the oversight, all the percentage signs have now been removed from the x-axis tick labels.

(10) The authors might show the (even not significant) results of the associations between changes in growth and cognitive flexibility in supplementary information.

As suggested, we added the table reporting the results of the associations between changes in growth (DWLZ) and cognitive flexibility in the supplementary materials (Table SI3). We additionally explored the relationship between changes in functional connectivity and cognitive flexibility as suggested by Reviewer #2 and we reported these additional analyses in the text as follows: “We also explored whether changes in growth and changes in functional connectivity between 5 and 24 months were associated with cognitive flexibility at preschool age, but we did not find any significant association (Table SI3 and Table SI4).” (line 213, results section).

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: Functional brain network organization, measured by functional connectivity (FC), reflects key neurodevelopmental processes for healthy development. Early exposure to adversity, e.g. undernutrition, affects neurodevelopment, observable via disrupted FC, and leads to poorer outcomes from preschool age onward. We assessed longitudinally the impact of early growth trajectories on developmental FC in a rural Gambian population from age 5 to 24 months. To investigate how these early trajectories relate to later childhood outcomes, we assessed cognitive flexibility at 3-5 years. We observed that early physical growth before the fifth month of life drove optimal developmental trajectories of FC that in turn predicted cognitive flexibility at pre-school age. In contrast to previously studied developmental populations, this Gambian sample exhibited long-range interhemispheric FC that decreased with age. Our results highlight the measurable effects that poor growth in early infancy has on brain development and the subsequent impact on pre-school age cognitive development, underscoring the need for early life interventions throughout global settings of adversity.

Revised version: Functional brain network organization, measured by functional connectivity (FC), reflects key neurodevelopmental processes for healthy development. Early exposure to adversity, e.g. undernutrition, affects neurodevelopment, observable via disrupted FC, and leads to poorer outcomes from preschool age onward. We assessed longitudinally the impact of early growth trajectories on developmental FC in a rural Gambian population from age 5 to 24 months. To investigate how these early trajectories relate to later childhood outcomes, we assessed cognitive flexibility at 3-5 years. We observed that early physical growth before the fifth month of life drove optimal developmental trajectories of FC that in turn predicted cognitive flexibility at pre-school age. In contrast to previously studied developmental populations, this Gambian sample exhibited long-range interhemispheric FC that decreased with age. Our results highlight the measurable effects that poor growth in early infancy has on brain development and the possible subsequent impact on pre-school age cognitive development, underscoring the need for early life interventions throughout global settings of adversity.

Introduction

The first 1000 days of life are of paramount importance for human brain (1) and body development (2, 3) Early adversity negatively impacts infant development, which leads to long-term consequences from delayed milestones in childhood development to lowered productivity at the individual and societal levels (4–6). Therefore, assessing trajectories of brain development in infants exposed to early adversity is essential to understand and mitigate effects on brain development and subsequent outcomes at school age and beyond (7, 8).

Early undernutrition engenders detrimental effects on a range of cognitive skills (9), with long-lasting effects through adulthood (10). Interventions using therapeutic feeding regimens and supplements have been shown to have a strong causal impact on childhood undernutrition, brain development, and subsequent neurodevelopmental outcomes (11–13). Despite numerous attempts of interventions, global rates of undernutrition remain high (14), with infants in low- and middle-income countries (LMICs) at greatest risk (15). Up to one-third of infants in LMICs are at risk of not meeting standard milestones of social, motor, linguistic, and cognitive development by pre-school age (16, 17), which can lead to lifelong consequences that impact social and economic development at individual, national, and global scales (18). Thus, it is essential to study neurodevelopment throughout the first two years in regions with high rates of undernutrition in order to optimize early interventions. To identify effective early interventions, we must first identify the mechanisms that link undernutrition to later cognitive outcomes. Effects of undernutrition on the developing brain may manifest through alterations in function brain networks (19), as shown in school-age children exposed to chronic poverty (20). However, our understanding of the relationships between early undernutrition, brain connectivity, and later cognitive outcomes is still very limited.

Functional brain network organization is commonly assessed with resting-state functional connectivity (FC) via temporal correlations in brain activity as measured with functional magnetic resonance imaging (fMRI) (21, 22). Based on data from fMRI, current models hypothesize that FC patterns mature throughout early development (23–27), where in typically developing brains, adult-like networks emerge over the first years of life as long-range functional connections between pre-frontal, parietal, temporal, and occipital regions become stronger and more selective (28–31). This maturation in FC has been shown to be related to the cascading maturation of myelination and synaptogenesis (32, 33) - fundamental processes for healthy brain development (34). Therefore, disrupted patterns of FC may reflect disrupted anatomical maturation of brain circuits and systems. For example, preterm infants with severe diffuse white matter impairment exhibited lower levels of FC in executive function networks compared with preterm infants with no or mild white matter impairment (35). Importantly, normative developmental patterns may be disrupted and even reversed in clinical conditions that impact development; e.g., increased short-range and reduced long-range FC have been observed in preterm infants (36) and in children with autism spectrum disorder (37, 38).

While widely used neuroimaging modalities such as fMRI can elucidate typical and atypical developmental trajectories, they are often poorly suited for lower resource neuroimaging contexts due to limited portability, need for specialized facilities, and the high cost associated with these methods (28, 36, 39). To address the need for neuroimaging studies in lower resource environments including (but not limited to) in LMICs, researchers have increasingly turned to more portable tools, including electroencephalography (EEG) (40), functional near-infrared spectroscopy (fNIRS) (41–43) and diffuse correlation spectroscopy (DCS) (44). These tools enable the creation of mobile neuroimaging laboratories that can be deployed virtually anywhere, eliminating practical constraints imposed by costly and immobile methods. Indeed, research based in The Gambia, Guinea-Bissau, Bangladesh and Ivory Coast have established feasibility of these tools to perform brain imaging studies outside of a specialized lab environment, with recent studies reporting altered cortical physiology related to early adversity and undernutrition (40, 43–45). For example, using EEG, Xie and colleagues showed that in 2-3 year old Bangladeshi children, growth faltering was associated with FC in the theta and beta frequency bands, which was negatively related to children’s IQ score at 4 years (46). With regards to fNIRS, it has a better spatial resolution and anatomical specificity than EEG, thus providing more precise and reliable localization of brain networks (47, 48). Moreover, fNIRS facilitates testing in awake and engaged infants, making the results more comparable with adult FC studies that are performed on awake participants (49). Given these unique strengths, fNIRS has been used to study FC in longitudinal studies in awake infants in high-resource settings (31), and the portability and the low cost of this equipment has fostered implementation outside conventional labs and in LMICs (50). Because fNIRS overcomes logistical (MRI) and resolution (EEG) challenges of other common modalities, we chose it for this study to assess early life brain health of children in LMIC exposed to early severe undernutrition.

Earlier version: The goal of the study was to investigate early physical growth in infancy, developmental trajectories of brain FC across the first two years of life, and cognitive outcome at school age in a longitudinal cohort of infants and children from rural Gambia, an environment with high rates of maternal and child undernutrition. Specifically, we aimed to: (i) investigate whether differences in physical growth through the first two years of life are related to FC at 24 months, and (ii) investigate if early FC has an impact on cognitive outcome at pre-school age in these children. Data were collected from N=204 children as part of the Brain Imaging for Global Health project (BRIGHT; globalfnirs.org/the-bright-project), a longitudinal study examining infant development from term birth to 24-months of age. We acquired longitudinal measurements of brain FC in awake Gambian infant participants at five age points (5, 8, 12, 18, and 24 months) using fNIRS as the infants watched calming videos. We also acquired anthropometric measures (i.e., weight and length) as indicators of nutritional status at birth, at 7-to-14 days, one month of age, and at all five fNIRS data acquisition time points. Additionally, because developmental delays interfere with everyday functioning starting within the pre-school age range (16), we assessed cognitive flexibility, a core executive function (51), at ages three and five years. We hypothesized that (i) long-range FC would increase and short-range FC would decrease throughout the first two years of life, (ii) that early positive physical growth would significantly predict long-range FC, and that (iii) long-range FC strength would predict cognitive flexibility performance in pre-schoolers. Our results provide novel insights into the developmental trajectory of the functional organization of the developing brain, its relationship with nutritional adversity exposure and status, and the subsequent effects on preschool-age cognitive outcomes.

Revised version: The goal of the study was to investigate early physical growth in infancy, developmental trajectories of brain FC across the first two years of life, and cognitive outcome at school age in a longitudinal cohort of infants and children from rural Gambia, an environment with high rates of maternal and child undernutrition. Specifically, we aimed to(i)investigate whether differences in physical growth through the first two years of life are related to FC at 24 months, and (ii) investigate if early FC has an impact on cognitive outcome at pre-school age in these children. Data were collected from N=204 children as part of the Brain Imaging for Global Health project (BRIGHT; http://globalfnirs.org/the-bright-project), a longitudinal study examining infant development from term birth to 24-months of age. We acquired longitudinal measurements of brain FC in awake Gambian infant participants at five age points (5, 8, 12, 18, and 24 months) using fNIRS as the infants watched calming videos. We also acquired anthropometric measures (i.e., weight and length) as indicators of nutritional status at birth, at 7-to-14 days, one month of age, and at all five fNIRS data acquisition time points. Additionally, because developmental delays interfere with everyday functioning starting within the pre-school age range (16), we assessed cognitive flexibility, a core executive function (51), at ages three and five years. We hypothesized that (i) long-range FC would increase and short-range FC would decrease throughout the first two years of life, (ii) that early positive physical growth would significantly predict long-range FC, and that (iii) long-range FC strength would predict cognitive flexibility performance in pre-schoolers. Our results provide novel insights into the developmental trajectory of the functional organization of the developing brain, its relationship with nutritional adversity exposure and status, and the subsequent effects on preschool-age cognitive outcomes.

Results

Gambian infants exhibit robust FC development throughout first two years

Earlier version: The repeated assessments of brain FC throughout the first two years of life allowed us to analyse developmental trajectories of the maturation of early brain networks and how they relate to measures of infants’ growth as well as to cognitive flexibility at preschool age (Figure 1A). To estimate how FC changed between 5-24 months of age in Gambian children, we performed linear mixed models (LMM) using data from the 132 infants with valid data for at least two time points. To increase statistical power while minimizing errors due to variability in head anatomy and cap placement, we divided the array into six sections and calculated functional connectivity (FC) between all possible 21 interhemispheric homotopic, intrahemispheric, fronto-posterior, and crossed connections (Figure SI1). Results on the Fisher-z transformed correlation coefficients (z-RHO) of the oxygenated haemoglobin (HbO2) showed that frontal interhemispheric homotopic FC decreased with age (F=11.03, p<0.001), and that left (F=5.8, p<0.001) and right fronto-middle (F=4.86, p<0.001) and right fronto-posterior (F=5.52, p<0.001) FC increased with age (Figure 2). Results on the z-RHO scores of the deoxygenated haemoglobin (HHb) largely agreed with the results on HbO2 (Table 1 and Figure 2). Frontal interhemispheric FC significantly decreased with age, showing positive correlations at five months, and negative values at 24 months (Figure 2B and Figure 2C). The FC between fronto-middle and frontal-posterior that significantly increased with age showed strong anticorrelation values at early ages that gradually decreased but were still negative at 24 months (Figure 2B and Figure 2C). Notably, results from LMMs performed on data pre-processed without the global signal regression (GSR) (see results in SI, Table SI1 and Figure SI2) confirmed a decreasing trajectory of the frontal interhemispheric homotopic FC and an increasing trajectory of bilateral short-range FC with age.

Revised version: The repeated assessments of brain FC throughout the first two years of life allowed us to analyse developmental trajectories of the maturation of early brain networks and how they relate to measures of infants’ growth as well as to cognitive flexibility at preschool age (Figure 1A). To estimate how FC changed between 5-24 months of age in Gambian children, we performed linear mixed models (LMM) using data from the 132 infants with valid data for at least two time points. To increase statistical power while minimizing errors due to variability in head anatomy and cap placement, we divided the array into six sections and calculated functional connectivity (FC) between all possible 21 interhemispheric homotopic, intrahemispheric within section, fronto-posterior, and crossed connections (Figure SI1). Results on the Fisher-z transformed correlation coefficients (z-RHO) of the oxygenated haemoglobin (HbO2) showed that frontal interhemispheric homotopic FC decreased with age (F=11.03, p<0.001), and that left (F=5.8, p<0.001) and right fronto-middle (F=4.86, p<0.001) and right fronto-posterior (F=5.52, p<0.001) FC increased with age (Figure 2).

Experimental design.

Earlier version: (A) Measures taken in the BRIGHT project used in this work. The measuring tape represents anthropometric measures, the brain represents fNIRS FC and the test represents the cognitive flexibility assessment. (B) Schematic representation of the spatial layout of the fNIRS array. Sources are marked with red stars, detectors are marked with blue circles, channels are marked with grey lines and numbered with black circles. The channels/optodes used as a reference for the tragus are highlighted in green. (C) Participants undergoing fNIRS testing at different ages. Image reproduced with permission from (41).

Revised version: (A) Measures taken in the BRIGHT project used in this work. The measuring tape represents anthropometric measures, the brain represents fNIRS FC and the test represents the cognitive flexibility assessment. (B) Schematic representation of the spatial layout of the fNIRS array. Sources are marked with red stars, detectors are marked with blue circles, channels are marked with grey lines and numbered with black circles. The channels/optodes used as a reference for the tragus are highlighted in green.

Linear mixed models results showing FC that displayed a statistically significant change with age.

Earlier version: (A) Significant results of the linear mixed model, blue indicates connections that decreased with age, red indicates connections that increase with age. (B) Mean and standard error of the mean (SE) of the functional connections that changed with age. Error bars are 1 SE. (C) Violin plot showing the mean ± standard deviation (SD) (red circles and lines) and the individual variability (coloured dots) of the FC that showed a change with time.

Revised version: (A) Significant results of the linear mixed model, blue indicates connections that decreased with age, red indicates connections that increase with age. (B) Mean and standard error of the mean (SE) of the functional connections that changed with age (HbO2). Error bars are 1 SE. (C) Violin plot showing the mean ± standard deviation (SD) (red circles and lines) and the individual variability (coloured dots) of the FC that showed a change with time (HbO2).

Added in the revised version: Results on the z-RHO scores of the deoxygenated haemoglobin (HHb) largely agreed with the results on HbO2 (Table 1 and Figure 2). Frontal interhemispheric FC significantly decreased with age, showing positive correlations at five months, and negative values at 24 months (Figure 2B and Figure 2C). The FC between fronto-middle and frontal-posterior that significantly increased with age showed strong anticorrelation values at early ages that gradually decreased but were still negative at 24 months (Figure 2B and Figure 2C).

Added in the revised version: Notably, results from LMMs performed on data pre-processed without the global signal regression (GSR) (see results in SI, Table SI1 and Figure SI2) confirmed a decreasing trajectory of the frontal interhemispheric homotopic FC and an increasing trajectory of bilateral short-range FC with age.

Early growth trajectories predict functional connectivity at 24 months

Earlier version: To investigate the impact of early nutritional status on FC at 24 months, we used multiple regression with the infant growth trajectory (delta weight for length z-score between all time points, ΔWLZ) and FC at 24 months, adjusting for WLZ at birth or head-circumference z-score (HCZ) at 7/14 days. To maximise power, we considered only those FC that showed a statistically significant change with age. The ΔWLZ between birth/1 month and older ages positively predicted frontal interhemispheric homotopic FC at 24 months (ΔWLZ birth with older ages all p<0.02 FDR corrected, ΔWLZ 1 month with older ages all p<0.03, see Figure SI3 for an example scatterplot) and negatively predicted left and right fronto-middle FC at 24 months (p<0.04). The ΔWLZ between birth and 1 month positively predicted right fronto-posterior FC at 24 months (p=0.006, FDR corrected) while ΔWLZ between 1 month and older ages negatively predicted right fronto-posterior FC at 24 months (all p<0.01, FDR corrected). Interestingly, ΔWLZ between 5/8/12/18 months and older ages did not show statistically significant impacts on any of the FC assessed (Table 2 and Figure 3A).

Revised version: To investigate the impact of early nutritional status on FC at 24 months, we used multiple regression with the infant growth trajectory (delta weight for length z-score between all time points, ΔWLZ) and FC at 24 months, adjusting for WLZ at birth or head-circumference z-score (HCZ) at 7/14 days. To maximise power, we considered only those FC that showed a statistically significant change with age. The ΔWLZ between birth/1 month and older ages positively predicted frontal interhemispheric homotopic FC at 24 months (ΔWLZ birth with older ages all p<0.02 FDR corrected, ΔWLZ 1 month with older ages all p<0.03, see Figure SI3 for some examples of scatterplots) and negatively predicted left and right fronto-middle FC at 24 months (p<0.04). The ΔWLZ between birth and 1 month positively predicted right fronto-posterior FC at 24 months (p=0.006, FDR corrected) while ΔWLZ between 1 month and older ages negatively predicted right fronto-posterior FC at 24 months (all p<0.01, FDR corrected). Interestingly, ΔWLZ between 5/8/12/18 months and older ages did not show statistically significant impacts on any of the FC assessed (Table 2 and Figure 3A).

Associations between FC, early growth and later cognitive flexibility.

Earlier version: (A) Significant positive associations are in green, significant negative associations are in orange, and non-significant associations are in blue. * indicates regressions that survived FDR correction for multiple comparisons. (B) Schematic representations of the early FC connections shown to predict cognitive flexibility in preschoolers. Significant positive associations are in green, significant negative associations are in orange.

Revised version: (A) Significant positive associations are in green, significant negative associations are in orange, and non-significant associations are in blue. * indicates regressions that survived FDR correction for multiple comparisons. (B) Schematic representations of the early FC connections shown to predict cognitive flexibility in preschoolers. Significant positive associations are in green, significant negative associations are in orange (these results did not survive corrections for multiple comparisons).

Early functional connectivity predicts cognitive flexibility at preschool age

Earlier version: To investigate whether FC early in life predicted cognitive flexibility at preschool age, we used multiple regression of FC across the first two years of life against later cognitive flexibility in preschoolers at three and five years. As per the analysis above, we focused on only those FC that showed a statistically significant change with age. Our results showed that frontal interhemispheric homotopic FC at 5 months (F(1,38)=4.21, p=0.047, R2=0.102), left frontal-middle (F(1,33)=7.86, p=0.009, R2=0.197) and right frontal-posterior FC at 12 months (F(1,38)=4.82, p=0.034, R2=0.115) positively predicted cognitive flexibility in young preschoolers. Right frontal-posterior FC at 8 months negatively predicted cognitive flexibility in young preschoolers (F(1,32)=5.6, p=0.024, R2=0.153). Frontal interhemispheric homotopic FC at 18 months (F(1,45)=4.82, p=0.030, R2=0.115), and right frontal-posterior FC at 24 months (F(1,49)=4.85, p=0.032, R2=0.092) positively predicted cognitive flexibility in older preschoolers. Left front-middle FC at 18 months negatively predicted cognitive flexibility in older preschoolers (F(1,45)=5.72, p=0.021, R2=0.115). While these associations between early functional connectivity and cognitive flexibility at preschool age show promise, none of these survived FDR correction for multiple comparisons (Figure 3B), and so should be considered as preliminary and useful for hypothesis generation for future studies.

Revised version: To investigate whether FC early in life predicted cognitive flexibility at preschool age, we used multiple regression of FC across the first two years of life against later cognitive flexibility in preschoolers at three and five years. As per the analysis above, we focused on only those FC that showed a statistically significant change with age. Our results showed that frontal interhemispheric homotopic FC at 5 months (F(1,38)=4.21, p=0.047, R2=0.102), left frontal-middle (F(1,33)=7.86, p=0.009, R2=0.197) and right frontal-posterior FC at 12 months (F(1,38)=4.82, p=0.034, R2=0.115) positively predicted cognitive flexibility in young preschoolers. Right frontal-posterior FC at 8 months negatively predicted cognitive flexibility in young preschoolers (F(1,32)=5.6, p=0.024, R2=0.153). Frontal interhemispheric homotopic FC at 18 months (F(1,45)=4.82, p=0.030, R2=0.115), and right frontal-posterior FC at 24 months (F(1,49)=4.85, p=0.032, R2=0.092) positively predicted cognitive flexibility in older preschoolers. Left front-middle FC at 18 months negatively predicted cognitive flexibility in older preschoolers (F(1,45)=5.72, p=0.021, R2=0.115). While these associations between early functional connectivity and cognitive flexibility at preschool age show promise, none of these survived FDR correction for multiple comparisons (Figure 3B and Table SI2), and so should be considered as preliminary and useful for hypothesis generation for future studies.

Added in the revised version: We also explored whether changes in growth and changes in functional connectivity between 5 and 24 months were associated with cognitive flexibility at preschool age, but we did not find any significant association (Table SI3 and Table SI4).

Discussion

This study investigated how early adversity via undernutrition drives brain functional connectivity throughout the first two years of life and how these early functional connections are associated with cognitive flexibility at preschool age. To capture the rapid neural development that takes place during early years in community LMIC settings we used fNIRS to assess brain connectivity (22, 24, 26, 31). Additionally, we recorded brain activity while infants were awake, thus making the results more comparable with those from adult studies. Our results show: (i) healthy growth trajectories during early infancy are crucial for healthy FC at 24 months; (ii) our cohort of Gambian infants exhibit atypical developmental trajectories of functional connectivity (FC) in over the first two years of life; and (iii) FC during the first two years of life may predict cognitive flexibility at preschool age.

Earlier version: First, we found that while left and right intra-hemispheric fronto-middle and right frontal-posterior FC increased with age, frontal inter-hemispheric FC decreased with age. Additionally, we observed long-range right frontal-posterior FC increased with age and exhibited positive correlations at 24 months, as consistent with previous infant studies (29, 31). Previous findings from typically developing infants measured in the US showed that long-range FC gradually strengthened with age, indicating maturing of functional networks that span distant brain regions (28, 29). Moreover, the presence of inter-hemispheric connections are known to indicate healthy development of FC (52) consistent with the continuous development of the corpus callosum from infancy until early adulthood (53).

Revised version: First, we found that while left and right intrahemispheric fronto-middle and right frontal-posterior FC increased with age, frontal inter-hemispheric FC decreased with age. Additionally, we observed long-range right frontal-posterior FC increased with age and exhibited positive correlations at 24 months, as consistent with previous infant studies (29, 31). Previous findings from typically developing infants measured in the US showed that long-range FC gradually strengthened with age, indicating maturing of functional networks that span distant brain regions (28, 29). Moreover, the presence of inter-hemispheric connections are known to indicate healthy development of FC (52) consistent with the continuous development of the corpus callosum from infancy until early adulthood (53).

Earlier version: Conversely, decreased inter-hemispheric FC has been observed where development can follow a more divergent and/or heterogeneous path, such as in autism (54). The observed decrease in frontal inter-hemispheric FC with increasing age may be due to the exposure to early life undernutrition adversity. We acknowledge that differences in FC could also be attributed to other environmental and cultural disparities between high-resource and low-resource settings, and future studies are needed to explore this further. Finally, our results indicated that the Gambian infant bilateral fronto-middle FC exhibited anticorrelation at five months, and then gradually became less negative with age. This pattern suggests that this fronto-middle FC may become in sync later after the second year of life. The fNIRS array covers regions over the front and middle cortical regions, putatively belonging to the default mode network and fronto-parietal network, which are normally anticorrelated (55), suggesting that the earlier observed anticorrelation pattern might indeed be adaptive. While further investigations are needed to confirm these findings, these results highlight a general disruption of the integration/segregation process of network development among the infants in our study (56).

Revised version: Conversely, decreased inter-hemispheric FC has been observed where development can follow a more divergent and/or heterogeneous path, such as in autism (54). The observed decrease in frontal inter-hemispheric FC with increasing age may be due to the exposure to early life undernutrition adversity. We acknowledge that differences in FC could also be attributed to other environmental and cultural disparities between high-resource and low-resource settings, and future studies are needed to explore this further. Finally, our results indicated that the Gambian infant bilateral fronto-middle FC exhibited anticorrelation at five months, and then gradually became less negative with age. This pattern suggests that this fronto-middle FC may become in sync later after the second year of life. The fNIRS array covers regions over the front and middle cortical regions, putatively belonging to the default mode network and fronto-parietal network, which are normally anticorrelated (55), suggesting that the earlier observed anticorrelation pattern might indeed be adaptive.

Added in the revised version: Importantly, these results remained significant even without GSR, indicating that our findings are not solely driven by preprocessing choices. While the use of GSR in FC studies remains debated (56), in the absence of short channels (which are difficult to use reliably with infants [57]) and external physiological measures, applying GSR represented the most appropriate preprocessing option. In fact, failure to correct for systemic physiological fluctuations can, in fact, lead to artificially elevated connectivity estimates in fNIRS data (58). While further investigations are needed to confirm these findings, these results highlight a general disruption of the integration/segregation process of network development among the infants in our study (59).

Earlier version: Second, a central goal of the current work was to investigate whether infant nutritional status, as measured by change in weight for length z-scores (ΔWLZ), was associated with functional connectivity at 24 months. First, we observed that ΔWLZ between neonatal time points and older ages positively predicted frontal interhemispheric homotopic FC at 24 months, which supports the idea that the decreasing interhemispheric connectivity observed with age might be atypical, as greater increases in ΔWLZ were found to be positively associated with stronger interhemispheric FC. This is consistent with research showing an impact of nutritional status on long-range connections and brain networks development (57–59) and the decreased long-range FC found in preterm babies with faltering growth (60–62). A common feature of the observed associations between ΔWLZ and FC is that they were statistically significant only when ΔWLZ was calculated including measures collected neonatally (i.e., at birth or one month of age). Changes in growth later in development (i.e. ΔWLZ between 5, 8, 12, or 18 months and older ages) did not show statistically significant impacts on any of the FC assessed. This provides support for the hypothesis that undernutrition during the first months of life is more impactful on brain development than undernutrition occurring later in infancy and early childhood. While the impact of undernutrition on brain development has been documented in LMICs (46), herein, we provided empirical evidence that growth faltering specifically in infants younger than five months of age impacts observable development of functional brain networks in the second year of life. This result also suggests that early undernutrition has a lasting impact on subsequent cognitive skills, even in children who show subsequent catch-up growth.

Revised version: Second, a central goal of the current work was to investigate whether infant nutritional status, as measured by change in weight for length z-scores (ΔWLZ), was associated with functional connectivity at 24 months. First, we observed that ΔWLZ between neonatal time points and older ages positively predicted frontal interhemispheric homotopic FC at 24 months, which supports the idea that the decreasing interhemispheric connectivity observed with age might be atypical, as greater increases in ΔWLZ were found to be positively associated with stronger interhemispheric FC. This is consistent with research showing an impact of nutritional status on long-range connections and brain networks development (60–62) and the decreased long-range FC found in preterm babies with faltering growth (63–65). A common feature of the observed associations between ΔWLZ and FC is that they were statistically significant only when ΔWLZ was calculated including measures collected neonatally (i.e., at birth or one month of age). Changes in growth later in development (i.e. ΔWLZ between 5, 8, 12, or 18 months and older ages) did not show statistically significant impacts on any of the FC assessed. This provides support for the hypothesis that undernutrition during the first months of life is more impactful on brain development than undernutrition occurring later in infancy and early childhood. While the impact of undernutrition on brain development has been documented in LMICs (46), herein, we provided empirical evidence that growth faltering specifically in infants younger than five months of age impacts observable development of functional brain networks in the second year of life. Future studies may be needed to pinpoint which specific brain networks are impacted. This result also suggests that early undernutrition has a lasting impact on subsequent cognitive skills, even in children who show subsequent catch-up growth.

Earlier version: The other primary aim of this study was to investigate whether early FC development was related to cognitive flexibility in preschoolers at three and five years of age. Long range FC (interhemispheric homotopic and frontal-posterior) positively predicted performance in the cognitive flexibility task both in younger and older preschoolers. This is consistent with the fact that the integration of distant brain regions fosters cognitive development, especially in prefrontal cortex (63, 64), possibly driven by maturation within the fronto-parietal network that has been widely documented as a primary neural underpinning of cognitive flexibility (65). Additionally, left fronto-middle FC at 12 months positively predicted cognitive flexibility in young preschoolers, but the same connection at 18 months negatively predicted cognitive flexibility in older preschoolers. This shift in direction of the relationship between FC and cognitive flexibility might indicate the shift between short-range and long-range FC during development, where short-range FC still promotes cognitive development early in life, while becoming less impactful as the child ages possibly because children develop a wider range of strategies to support cognitive flexibility as they age. A recent review on the neural underpinnings of cognitive flexibility highlighted that right-lateralized activations underlying this skill increased with age, while activations over the left hemisphere decreased with age (51). Our findings are consistent with this view, suggesting an increase in functional specialisation and segregation with development. While our results are consistent with previous studies, we acknowledge that the significant association between early FC and later cognitive flexibility does not withstand multiple comparisons. Therefore, we encourage future studies that may replicate these findings with a larger sample.

Revised version: The other primary aim of this study was to investigate whether early FC development was related to cognitive flexibility in preschoolers at three and five years of age. Long range FC (interhemispheric homotopic and frontal-posterior) positively predicted performance in the cognitive flexibility task both in younger and older preschoolers. This is consistent with the fact that the integration of distant brain regions fosters cognitive development, especially in prefrontal cortex (66, 67), possibly driven by maturation within the fronto-parietal network that has been widely documented as a primary neural underpinning of cognitive flexibility (68). Additionally, left fronto-middle FC at 12 months positively predicted cognitive flexibility in young preschoolers, but the same connection at 18 months negatively predicted cognitive flexibility in older preschoolers. This shift in direction of the relationship between FC and cognitive flexibility might indicate the shift between short-range and long-range FC during development, where short-range FC still promotes cognitive development early in life, while becoming less impactful as the child ages possibly because children develop a wider range of strategies to support cognitive flexibility as they age. A recent review on the neural underpinnings of cognitive flexibility highlighted that right-lateralized activations underlying this skill increased with age, while activations over the left hemisphere decreased with age (51). Our findings are consistent with this view, suggesting an increase in functional specialisation and segregation with development. While our results are consistent with previous studies, we acknowledge that the significant associations between early FC and later cognitive flexibility do not withstand multiple comparisons. Therefore, we encourage future studies that may replicate these findings with a larger sample.

Earlier version: Some limitations of the current study merit discussion. Although the BRIGHT sample is one of the largest longitudinal neuroimaging studies in LMIC to assess longitudinally this many age points during the first two years of life, the sample size may still be underpowered to assess the modest effects of the regression models. Second, additional factors may affect FC that were not considered in our models. Applying more advanced statistical modelling methods and structural equation modelling analyses may provide greater insight with further investigations in contexts of adversity and, in turn, establish which outcomes are predicted by FC. That said, results from the regression models were statistically robust, as most survived FDR correction for multiple comparisons, and were significant after including HCZ at 7-14 days and WLZ at birth in the models. It is also worth noting that in line with previous FC fNIRS studies (31, 66), linear mixed model results on the HbO2 and the HHb signals showed consistency in the FC changes with age measured in the two chromophores, and similar trajectories when removing the global signal regression step from the pre-processing, suggesting that the results were reliable and statistically robust. Future collaborations between projects studying the impact of adversity on brain development in both high resource and LMIC settings are crucial to advance the field and therefore inform efficient strategies – including the timing - of intervention. Even given the large sample in our study, we were underpowered to test for group comparisons between sets of infants with distinct undernutrition growth profiles, e.g., infants with early poor growth that later resolved and infants with standard growth early that had a poor growth later. We were also underpowered to test the associations between early growth and FC on clinically undernourished infants (defined as having ΔWLZ two standard deviations below the mean). However, our findings provide a solid foundation for future research on the impacts of undernutrition during the first months of life on later brain health and cognitive abilities throughout childhood.

Revised version: Some limitations of the current study merit discussion. Although the BRIGHT sample is one of the largest longitudinal neuroimaging studies in LMIC to assess longitudinally this many age points during the first two years of life, the sample size may still be underpowered to assess the modest effects of the regression models. Second, additional factors may affect FC that were not considered in our models. Applying more advanced statistical modelling methods and structural equation modelling analyses may provide greater insight with further investigations in contexts of adversity and, in turn, establish which outcomes are predicted by FC. That said, results from the regression models were statistically robust, as most survived FDR correction for multiple comparisons, and were significant after including HCZ at 7-14 days and WLZ at birth in the models. It is also worth noting that in line with previous FC fNIRS studies (31, 69), linear mixed model results on the HbO2 and the HHb signals showed consistency in the FC changes with age measured in the two chromophores, and similar trajectories when removing the global signal regression step from the pre-processing, suggesting that the results were reliable and statistically robust. Future collaborations between projects studying the impact of adversity on brain development in both high resource and LMIC settings are crucial to advance the field and therefore inform efficient strategies – including the timing - of intervention. Even given the large sample in our study, we were underpowered to test for group comparisons between sets of infants with distinct undernutrition growth profiles, e.g., infants with early poor growth that later resolved and infants with standard growth early that had a poor growth later. We were also underpowered to test the associations between early growth and FC on clinically undernourished infants (defined as having ΔWLZ two standard deviations below the mean). However, our findings provide a solid foundation for future research on the impacts of undernutrition during the first months of life on later brain health and cognitive abilities throughout childhood.

Earlier version: In conclusion, herein we showed that, as expected, positive growth trajectories during the first few months of life were positively associated with stronger interhemispheric frontal FC and negatively associated with bilateral short-range FC. Interestingly, while Gambian infants had expected increasing left and right intra-hemispheric fronto-middle and right frontal-posterior FC with age, frontal inter-hemispheric FC decreased with age, which is inconsistent with extant data from other populations studied in Europe and US. Of particular note, growth during the first few months after birth showed a significant impact on brain health, as measured with FC, while later growth changes were not related to variability in FC. Moreover, long-range FC predicted the performance on a cognitive flexibility task at preschool age. Taken together, our results show the impact that early growth has on functional brain development, which in turn affects cognitive flexibility in preschoolers. Finally, this study emphasizes the importance of the timing of interventions, for which benefits are visible at preschool age and beyond.

Revised version: In conclusion, herein we showed that, as expected, positive growth trajectories during the first few months of life were positively associated with stronger interhemispheric frontal FC and negatively associated with bilateral short-range FC. Interestingly, while Gambian infants had expected increasing left and right intrahemispheric fronto-middle and right frontal-posterior FC with age, frontal inter-hemispheric FC decreased with age, which is inconsistent with extant data from other populations studied in Europe and US. Of particular note, growth during the first few months after birth showed a significant impact on brain health, as measured with FC, while later growth changes were not related to variability in FC. Moreover, long-range FC predicted the performance on a cognitive flexibility task at preschool age. Taken together, our results show the impact that early growth has on functional brain development, which in turn affects cognitive flexibility in preschoolers.

Added in the revised version: Finally, this study emphasizes the importance of the timing of interventions, for which benefits are visible at preschool age and beyond.

Materials and Methods

Experimental Design

Our study investigated the developmental trajectory of brain health as assessed with functional connectivity (FC) in Gambian infants and its relationship with early growth, as a proxy for nutritional status, and later cognitive outcomes. To do, so we estimated FC from data acquired with fNIRS longitudinally at five time points between age five and 24 months in the context of the BRIGHT Project. The BRIGHT project sought to provide brain function-for-age curves from the UK and Gambian infants and to gain insight into the effects that issues related to living in low-resource settings may have on brain development. We then tested for relationships between brain FC at age 24 months with measures of early growth, as indexed by changes in weight-for-length z-scores (reflecting body weight in proportion to attained growth in length) at one month of age, and at each of the four subsequent visits (details provided below). We also tested for associations between brain FC at each time point with later cognitive outcome, as indexed by cognitive flexibility assessed in three- and five-year-olds preschool-age children (Figure 1A). Ethical approval was granted by the joint Gambia Government - MRC Unit the Gambia Ethics Committee, and written informed consent was obtained from all parents prior to participation.

Participants

Earlier version: We used data from N=204 infants recruited in rural community settings of The Gambia as part of the BRIGHT project. The study was a prospective, longitudinal study with fNIRS assessments planned for awake infants at ages 5, 8, 12, 18 and 24 months. All infants were born full term (37–42 weeks gestation). fNIRS data were additionally acquired from 1-month-old infants while asleep. However, as sleep stages may impact connectivity measures (98), these data are not part of the current analysis. Reasons for exclusion from the fNIRS analysis were: i) fussiness; ii) experimental errors (missing pictures of the headgear placement or any other technical issue); iii) poor headgear placement (see supplementary information for Assessment of headgear placement); iv) low-quality data (more than 40% of the channels excluded as previously done (67–70)). Table SI2 summarises the details of the included and excluded participants for each visit, showing an attrition rate within the standard range for infant fNIRS studies (71) (see Table 3 for the demographic information of the participants included in the analyses at each age).

Revised version: We used data from N=204 infants recruited in rural community settings of The Gambia as part of the BRIGHT project. The study was a prospective, longitudinal study with fNIRS assessments planned for awake infants at ages 5, 8, 12, 18 and 24 months. All infants were born full term (37–42 weeks gestation). fNIRS data were additionally acquired from 1-month-old infants while asleep. However, as sleep stages may impact connectivity measures (98), these data are not part of the current analysis. Reasons for exclusion from the fNIRS analysis were: i) fussiness; ii) experimental errors (missing pictures of the headgear placement or any other technical issue); iii) poor headgear placement (see supplementary information for Assessment of headgear placement); iv) low-quality data (more than 40% of the channels excluded as previously done [70–73]). Table SI5 summarises the details of the included and excluded participants for each visit, showing an attrition rate within the standard range for infant fNIRS studies (74) (see Table 3 for the demographic information of the participants included in the analyses at each age).

Anthropometric measures

Measurements of length, weight, and head circumference (HC) were collected on all infants at birth, at 7/14 days, 1 month, and in all visits from 5 to 24 months, by trained field workers using calibrated tools. Length was measured using a Harpenden Infantometer length board (Holtain Ltd) to a precision of 0.1 cm. Weight was obtained using an electronic baby scale (Model 336, SECA) to a precision of 0.01 kg. Finally, HC was measured around the maximum circumference of the head (forehead to occiput) using stretch-proof measuring tape (Model 201, SECA) to the nearest 0.1 cm. Each measure was taken in triplicate and the mean of the three measures was used in analyses.

fNIRS data acquisition

Earlier version: fNIRS data were collected using the NTS topography system (Gowerlabs Ltd., London, UK). This system uses two continuous wavelengths of near-infrared light (780 and 850 nm) to detect changes in HbO2 and HHb concentrations, using a sampling rate of 10 Hz (72).

Revised version: fNIRS data were collected using the NTS topography system (Gowerlabs Ltd., London, UK). This system uses two continuous wavelengths of near-infrared light (780 and 850 nm) to detect changes in HbO2 and HHb concentrations, using a sampling rate of 10 Hz (75).

Earlier version: Infants were assessed using a custom-built fNIRS headgear consisting of two arrays, including 14 sources and 12 detectors to create a total of 34 channels, covering bilateral frontal, inferior-frontal and temporal regions, with a source-detector separation of 2 cm (Figure 1B and Figure 1C). Infants underwent several fNIRS tasks as part of the BRIGHT project, focusing on social processing, deferred imitation, and habituation response. Therefore, brain regions for data acquisition were chosen to maximize the likelihood of recording meaningful data for all tasks (73).

Revised version: Infants were assessed using a custom-built fNIRS headgear consisting of two arrays, including 14 sources and 12 detectors to create a total of 34 channels, covering bilateral frontal, inferior-frontal and temporal regions, with a source-detector separation of 2 cm (Figure 1B and Figure 1C). Infants underwent several fNIRS tasks as part of the BRIGHT project, focusing on social processing, deferred imitation, and habituation response.

Added in the revised version: Therefore, brain regions for data acquisition were chosen to maximize the likelihood of recording meaningful data for all tasks (76).

Earlier version: For each infant, head measurements were taken to enable the alignment of the headgear with the 10–20 placement coordinates (74). Over the front of the head, the headgear was placed so that a vertical red line marking the centre of the band aligned with the nasion with the silicon band laying just above, and in line with, the eyebrows. The headgear was placed so that the reference optode (the third lower optode from the back) was placed over the tragus.

Revised version: For each infant, head measurements were taken to enable the alignment of the headgear with the 10–20 placement coordinates (77). Over the front of the head, the headgear was placed so that a vertical red line marking the centre of the band aligned with the nasion with the silicon band laying just above, and in line with, the eyebrows. The headgear was placed so that the reference optode (the third lower optode from the back) was placed over the tragus.

To model for the headgear placement offline, photographs of the participants (one from the front and one from either side) were taken after the headgear was placed and again at the end of the fNIRS session. To assess the headgear placement, photographs of the participants had to be sufficiently good in quality (i.e. not blurred) and with the camera pointing at the ear, at a position directly perpendicular to the infant’s anatomical reference on the side of the head being photographed.

Infants were excluded from further analysis if the band was excessively high over the front above the eyebrows, by checking the displacement of the lower edge of the band with respect of the eyebrows with a “traffic-light” code (red, orange, green). Three experimenters independently classified the infants’ frontal headgear placement, and discrepancies were discussed until agreement was reached. Channels of those infants with valid placement over the front but horizontal lateral displacement between the ideal and the actual position of the reference (optode marked in green in Figure 1B) equal or greater than 1.6 cm were renumbered, so that each channel was shifted either backward or forward one full channel location in space

At each visit, the fNIRS data acquisition took place while infants sat on their parent’s lap. The parent was instructed to refrain from interacting with the infant during the stimuli presentation. Videos were of male and female Gambian adults singing nursery rhymes in Mandinka (for 70 seconds) and videos of toys in action (for 60 seconds), video blocks alternated, with each type shown to the infant twice, with the aim of keeping the infant awake, calm, attentive and as still as possible. The functional connectivity (FC) paradigm formed part of a greater battery of fNIRS paradigms that ran in a continuous and consistent order across all infants. The fNIRS FC acquisition lasted for a maximum of 520 seconds (two rounds of 260 seconds each) or were stopped as soon as the participant became fussy. Stimuli were presented using a MATLAB® custom-written stimulus presentation framework, Task Engine (sites.google.com/site/taskenginedoc) and Psychtoolbox on an Apple Macintosh computer.

fNIRS preprocessing and data-analysis

Earlier version: Data analyses were carried out using in-house programmes developed in MATLAB (MathWorks, Natick, MA) (Figure SI4 summarises fNIRS data preprocessing steps). Stringent multi-level data-quality assessments were used to ensure all data included in the group were of adequate quality. First, low-quality channels based on physiological indicators of quality were pruned using QT-NIRS (https://github.com/lpollonini/qt-nirs). The software takes advantage of the fact that the cardiac pulsation is also recorded in addition to haemodynamic activity associated with brain-driven variance including task-based responses and those contributing to functional connectivity (75), and quantifies its strength in the spectral and temporal domains with two measures: the Scalp Coupling Index (SCI) and the Peak Spectral Power (PSP) (for more details, see (76). For each infant, data quality was assessed channel-by-channel, and SCI and PSP were calculated with sliding 3-second windows (threshold SCI=0.70, threshold PSP=0.1, empirically defined and based on (77)) (76). Channels that had both SCI and PSP below threshold for more than 70% of the windows (76), were excluded from further pre-processing steps.

Revised version: Data analyses were carried out using in-house programmes developed in MATLAB (MathWorks, Natick, MA) (Figure SI4 summarises fNIRS data preprocessing steps). Stringent multi-level data-quality assessments were used to ensure all data included in the group were of adequate quality. First, low-quality channels based on physiological indicators of quality were pruned using QT-NIRS (https://github.com/lpollonini/qt-nirs). The software takes advantage of the fact that the cardiac pulsation is also recorded in addition to haemodynamic activity associated with brain-driven variance including task-based responses and those contributing to functional connectivity (78), and quantifies its strength in the spectral and temporal domains with two measures: the Scalp Coupling Index (SCI) and the Peak Spectral Power (PSP) (for more details, see (79). For each infant, data quality was assessed channel-by-channel, and SCI and PSP were calculated with sliding 3-second windows (threshold SCI=0.70, threshold PSP=0.1, empirically defined and based on [80]) (79). Channels that had both SCI and PSP below threshold for more than 70% of the windows (79), were excluded from further pre-processing steps.

Earlier version: Raw intensity data were converted to optical density via a logarithm after dividing by the temporal mean of each channel (hmrIntensity2OD.m function from Homer2 tool (78)) and bandpass filtered (0.009 – 3 Hz) (hmrBandpassFilt.m function from Homer2 (78)). Using first a wide filter leaves in the signal contamination from the high bandwidth information (i.e. respiration, pulse, etc.) that we aim to regress out next. To help manage the systemic arterial signal that is well known to contaminate fNIRS signals (79), we regressed out the mean of the signals across the array, a method also known as global signal regression (GSR) (80, 81) (Results of analysis performed also without GSR are reported in the Supplementary Materials). Hereafter, optical density data underwent a second bandpass filtered (0.009 – 0.08 Hz), as previously described (47). Motion artefacts were then rejected using global variance of temporal derivatives (GVTD, STD=5) in the NeuroDOT toolbox (https://www.nitrc.org/projects/neurodot/) (82) (Figure SI5 show how we tested the effect of different STD thresholds of GVTD on data inclusion). To further manage potential motion contamination of the data, we additionally removed five seconds before and after each motion artefact, and only chunks of data that were at least 20 seconds long were retained.

Revised version: Raw intensity data were converted to optical density via a logarithm after dividing by the temporal mean of each channel (hmrIntensity2OD.m function from Homer2 tool [81]) and bandpass filtered (0.009 – 3 Hz) (hmrBandpassFilt.m function from Homer2 [81]). Using first a wide filter leaves in the signal contamination from the high bandwidth information (i.e. respiration, pulse, etc.) that we aim to regress out next. To help manage the systemic arterial signal that is well known to contaminate fNIRS signals (82), we regressed out the mean of the signals across the array, a method also known as global signal regression (GSR) (83, 84) (Results of analysis performed also without GSR are reported in the Supplementary Materials). Hereafter, optical density data underwent a second bandpass filtered (0.009 – 0.08 Hz), as previously described (47). Motion artefacts were then rejected using global variance of temporal derivatives (GVTD, STD=5) in the NeuroDOT toolbox (https://www.nitrc.org/projects/neurodot/) (85) (Figure SI5 show how we tested the effect of different STD thresholds of GVTD on data inclusion). To further manage potential motion contamination of the data, we additionally removed five seconds before and after each motion artefact, and only chunks of data that were at least 20 seconds long were retained.

Earlier version: Optical density data were converted to relative concentrations of haemoglobin using the modified Beer–Lambert law (hmrOD2Conc.m function from Homer2 (78)). Differential path length factors were calculated (83) and adapted to wavelengths (780 and 850 nm) and ages (5 months = 5.24 and 4.25; 8 months = 5.26 and 4.26; 12 months = 5.27 and 4.28 ;18 months = 5.30 and 4.30; 24 months = 5.32 and 4.32).

Revised version: Optical density data were converted to relative concentrations of haemoglobin using the modified Beer–Lambert law (hmrOD2Conc.m function from Homer2 [81]). Differential path length factors were calculated (86) and adapted to wavelengths (780 and 850 nm) and ages (5 months = 5.24 and 4.25; 8 months = 5.26 and 4.26; 12 months = 5.27 and 4.28;18 months = 5.30 and 4.30; 24 months = 5.32 and 4.32).

Earlier version: To increase the statistical power of our analysis and to reduce the number of multiple comparisons for all statistical analysis to investigate developmental changes in connectivity, we averaged the concentration changes of the channels that survived pre-processing into sections. The sections were established via the 17 channels of each hemisphere which were grouped into front, middle and back (for a total of six regions) based on a previous co-registration of the BRIGHT fNIRS arrays onto age-appropriate templates (68) (Figure SI1). Each section represented a reasonable estimate of an anatomically consistent array coverage across the group, and was used to increase statistical power while minimizing errors due to variability in head anatomy and cap placement. For each participant, the Pearson-r correlation matrix between all the sections was calculated for both HbO2 and HHb, resulting in a 6×6 matrix of section-pair correlations. We then applied a Fisher z-transformation on the correlation matrix for further statistical analyses. Infants with at least 250 seconds of valid data after pre-processing were considered for further analyses. To choose the optimal minimum amount of valid data, per each infant at each time point for both the HbO2 and HHb signal, we correlated the connectivity matrices estimated in the first portion of data with the one in the last portion of data by increasing the seconds of data considered (i.e. the first 60 seconds with the last 60 seconds, the first 100 seconds with the last 100 seconds, the first 120 seconds with the last 120 seconds, etc.). We found that correlating the first and the last 250 seconds of valid data after pre-processing provided the highest percentage of infants with strong correlation between the first and the last portion of data, suggesting that FC patterns reached a usable stability after 250 seconds (Figure SI6).

Revised version: To increase the statistical power of our analysis and to reduce the number of multiple comparisons for all statistical analysis to investigate developmental changes in connectivity, we averaged the concentration changes of the channels that survived pre-processing into sections. The sections were established via the 17 channels of each hemisphere which were grouped into front, middle and back (for a total of six regions) based on a previous co-registration of the BRIGHT fNIRS arrays onto age-appropriate templates (71) (Figure SI1). Each section represented a reasonable estimate of an anatomically consistent array coverage across the group, and was used to increase statistical power while minimizing errors due to variability in head anatomy and cap placement. For each participant, the Pearson-r correlation matrix between all the sections was calculated for both HbO2 and HHb, resulting in a 6×6 matrix of section-pair correlations. We then applied a Fisher z-transformation on the correlation matrix for further statistical analyses. Infants with at least 250 seconds of valid data after pre-processing were considered for further analyses. To choose the optimal minimum amount of valid data, per each infant at each time point for both the HbO2 and HHb signal, we correlated the connectivity matrices estimated in the first portion of data with the one in the last portion of data by increasing the seconds of data considered (i.e. the first 60 seconds with the last 60 seconds, the first 100 seconds with the last 100 seconds, the first 120 seconds with the last 120 seconds, etc.). We found that correlating the first and the last 250 seconds of valid data after pre-processing provided the highest percentage of infants with strong correlation between the first and the last portion of data, suggesting that FC patterns reached a usable stability after 250 seconds (Figure SI6).

Cognitive flexibility at preschool age

Earlier version: Gambian infants from BRIGHT were cross-sectionally assessed at the age of 3 or 5 years for cognitive flexibility using the “card sorting” task from the tablet-based Early Years Toolbox (84) (http://www.eytoolbox.com.au). In this task, children were asked to sort cards (i.e. red rabbits or blue boats) either by colour (red or blue) or shape (rabbits or boats), and to flexibly switch back and forth between these rules (85).

Revised version: Gambian infants from BRIGHT were cross-sectionally assessed at the age of 3 or 5 years for cognitive flexibility using the “card sorting” task from the tablet-based Early Years Toolbox (87) (http://www.eytoolbox.com.au). In this task, children were asked to sort cards (i.e. red rabbits or blue boats) either by colour (red or blue) or shape (rabbits or boats), and to flexibly switch back and forth between these rules (88).

Earlier version: Due to disruptions to field work as a result of a political crisis in December 2016 - January 2017 that forced us to pause recruitment, data were collected in two ranges of preschool ages, younger preschoolers (age mean ± SD=47.96 ± 2.77 months, N=77) and older preschoolers (age mean ± SD=57.58 ± 2.11 months, N=84) (73).

Revised version: Due to disruptions to field work as a result of a political crisis in December 2016 - January 2017 that forced us to pause recruitment, data were collected in two ranges of preschool ages, younger preschoolers (age mean ± SD=47.96 ± 2.77 months, N=77) and older preschoolers (age mean ± SD=57.58 ± 2.11 months, N=84) (76).

Statistical Analyses

Removed in the revised version: To investigate developmental trajectories of functional connectivity between 5 and 24 months, we ran linear mixed models (86) which is the standard statistical test to analyse repeated measures and longitudinal data with dropouts (87). Statistical analysis was performed with SPSS Version 28.0 statistic software package. Compared to repeated measures ANOVA, linear mixed models account for within-participant dependence and allow for missing data by using only information from the individual at the other visits (88). All possible interhemispheric homotopic, intrahemispheric, fronto-posterior, and crossed connections between the six regions were inserted as dependent variables in a linear mixed model, for a total of 21 linear mixed models. Functional connectivity was modelled as the dependent variable as a function of age with a random participant effect and random errors.

Earlier version: Intercept and age were fixed effects, while within-participant dependence was modelled as a random effect (89). This same procedure was used in other longitudinal studies that explored developmental changes over time (31, 90, 91). The linear mixed model included the 132 infants who had valid data for at least two visits (55 infants contributed with data from two visits, 51 infants contributed with data from three visits, 22 infants contributed with data from four visits, four infants contributed with data from all of the five visits). The type of covariance between the observations was specified as Autoregression (AR) as two measures close in time of the same participant are likely to be correlated (92). To ensure statistical reliability, significant results from the linear mixed models were corrected for multiple comparisons using Bonferroni correction (p=0.00238).

Revised version: To investigate developmental trajectories of functional connectivity between 5 and 24 months, we ran linear mixed models (89) which is the standard statistical test to analyse repeated measures and longitudinal data with dropouts (90). Statistical analysis was performed with SPSS Version 28.0 statistic software package. Compared to repeated measures ANOVA, linear mixed models account for within-participant dependence and allow for missing data by using only information from the individual at the other visits (91). All possible interhemispheric homotopic, intrahemispheric within section, fronto-posterior, and crossed connections between the six regions were inserted as dependent variables in a linear mixed model, for a total of 21 linear mixed models. Functional connectivity was modelled as the dependent variable as a function of age with a random participant effect and random errors. Intercept and age were fixed effects, while within-participant dependence was modelled as a random effect (92). This same procedure was used in other longitudinal studies that explored developmental changes over time (31, 93, 94). The linear mixed model included the 132 infants who had valid data for at least two visits (55 infants contributed with data from two visits, 51 infants contributed with data from three visits, 22 infants contributed with data from four visits, four infants contributed with data from all of the five visits). The type of covariance between the observations was specified as Autoregression (AR) as two measures close in time of the same participant are likely to be correlated (95). To ensure statistical reliability, significant results from the linear mixed models were corrected for multiple comparisons using Bonferroni correction (p=0.00238).

Earlier version: Anthropometric measures were converted to age and sex adjusted z-scores that are based on World Health Organization Child Growth Standards (93). Weight-for-Length (WLZ) and Head Circumference (HCZ) z-scores were computed. Delta WLZ (ΔWLZ) were calculated for various periods (0-1 months, 0-5 months, etc.) by subtracting the z-score at the later age from that at the previous age. The use of delta z-scores as outcome measures allowed us to assess the impact of positive or negative deviation from the expected growth trajectories. To investigate the impact of undernutrition on FC development, we used ΔWLZ as independent variables in regression analyses on HbO2 (as the chromophore with the highest signal-to-noise ratio) FC at 24 months, our final time point of data collection. To maximise power, we considered only those FC that showed a statistically significant change with age. These analyses were adjusted by WLZ at birth and HCZ at 7/14 days, to more confidently assume that the associations between growth and FC were specific to the impact of change in WLZ postnatally and not confounded by the size or maturity of the infant at birth. We used HCZ measured at 7/14 days and not HCZ at birth as HC measures at birth are unreliable following vaginal delivery. Therefore, HCZ at 7/14 days is a more accurate head measure than HCAZ at birth. To ensure statistical reliability, results from the regression analyses on each FC were corrected for multiple comparisons using false discovery rate (FDR)(94) per each connection investigated, i.e. 21 possible ΔWLZ values per each connection.

Revised version: Anthropometric measures were converted to age and sex adjusted z-scores that are based on World Health Organization Child Growth Standards (96). Weight-for-Length (WLZ) and Head Circumference (HCZ) z-scores were computed. Delta WLZ (ΔWLZ) were calculated for various periods (0-1 months, 0-5 months, etc.) by subtracting the z-score at the later age from that at the previous age. The use of delta z-scores as outcome measures allowed us to assess the impact of positive or negative deviation from the expected growth trajectories. To investigate the impact of undernutrition on FC development, we used ΔWLZ as independent variables in regression analyses on HbO2 (as the chromophore with the highest signal-to-noise ratio) FC at 24 months, our final time point of data collection. To maximise power, we considered only those FC that showed a statistically significant change with age. These analyses were adjusted by WLZ at birth and HCZ at 7/14 days, to more confidently assume that the associations between growth and FC were specific to the impact of change in WLZ postnatally and not confounded by the size or maturity of the infant at birth. We used HCZ measured at 7/14 days and not HCZ at birth as HC measures at birth are unreliable following vaginal delivery. Therefore, HCZ at 7/14 days is a more accurate head measure than HCAZ at birth. To ensure statistical reliability, results from the regression analyses on each FC were corrected for multiple comparisons using false discovery rate (FDR)(97) per each connection investigated, i.e. 21 possible ΔWLZ values per each connection.

Earlier version: To investigate whether FC early in life predicted cognitive flexibility at preschool age, we regressed FC that showed a significant change across the first two years of life against later cognitive flexibility.

Revised version: To investigate whether FC early in life predicted cognitive flexibility at preschool age, we regressed later cognitive flexibility against FC that showed a significant change across the first two years of life.